<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 213819  -->
  <question type="stack">
    <name>
      <text>Détermination de la moyenne, maximum et minimum</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Compute the maximum of the entries in <a href="[[textdownload name="data.csv"]]{@stack_csv_formatter(data_set,headers)@}[[/textdownload]]">this file</a>.</p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The maximum is {@ta_1@}.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[
headers: ["Numbers"];
a: rand(10);
b: rand(10);
data_set: [[a], [b]];

/* Add headers row to data set */
data_set_hd: append(
    [headers],
    data_set
);

ta_1: max(a, b);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Max of {@a@} and {@b@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p></p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p><i class="fa fa-check"></i> Réponse correcte, bravo.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><i class="fa fa-adjust"></i> Votre réponse est partiellement correcte.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><i class="fa fa-times"></i> Réponse incorrecte.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Nice.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1355139652</deployedseed>
    <deployedseed>1475793692</deployedseed>
    <deployedseed>1405099942</deployedseed>
    <deployedseed>727222688</deployedseed>
    <deployedseed>1764488602</deployedseed>
    <deployedseed>1053645696</deployedseed>
    <deployedseed>1288483347</deployedseed>
    <deployedseed>475981690</deployedseed>
    <deployedseed>481921617</deployedseed>
    <deployedseed>1353148035</deployedseed>
    <deployedseed>866212287</deployedseed>
    <deployedseed>436996953</deployedseed>
    <deployedseed>1660075351</deployedseed>
    <deployedseed>730573743</deployedseed>
    <deployedseed>904169792</deployedseed>
    <deployedseed>1115876400</deployedseed>
    <deployedseed>670670297</deployedseed>
    <deployedseed>602749870</deployedseed>
    <deployedseed>404838629</deployedseed>
    <deployedseed>623713637</deployedseed>
    <deployedseed>2118629351</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>